# Intro to programming. Final exam

Canvas quiz + tasks in the notebook. Quiz contains questions on the results of tasks 1 and 2 from the notebook.

Input data is located in folder `inputs`. The example outputs are shown after the text of the task. You can also check how the outputs should look like in the folder `example_outputs`.

- Save outputs to the current directory, i.e. specify just the name of the file.
- In case if asserts fail, but the outputs look identical to you, leave it as is.
- No cheating, no blind copy-pasting from Stackoverflow.
- You are encouraged to use all programming super powers you now have✨

<img src='https://boyter.org/static/books/Image-uploaded-from-iOS.jpg' width='250px' align='center'>

```python
import pandas as pd 
import numpy as np

pd.set_option('display.max_columns', 30)
np.set_printoptions(suppress=True)

import matplotlib.pyplot as plt
import seaborn as sns

import pickle
t1 = pickle.load(open("example_outputs/t1.pk", "rb")) 
t4 = pickle.load(open("example_outputs/t4.pk", "rb")) 
t6 = pickle.load(open("example_outputs/t6.pk", "rb")) 
t7 = pickle.load(open("example_outputs/t7.pk", "rb")) 
t8 = pickle.load(open("example_outputs/t8.pk", "rb")) 
```

## Task 1. Thermodynamic data

1 point.

Read the data on protein stability changes upon mutations `inputs/1.csv`. Perform the following processing (leave/filter means select):

- Leave the entries for which the data on melting temperature changes (dTm) is available.
- Filter proteins with the length less than 300.
- Convert temperature to degrees Celsius and store them in the column 'T'.
- Create a new column 'MUTATION TYPE' with annotation of mutations by type (single or multipe). P.S. in the table multiple mutations are listed by comma.
- Leave the following columns (preserve the order): 'ID', 'PROTEIN', 'SOURCE', 'UNIPROT', 'MUTATION', 'MUTATION TYPE', 'dTm','T', 'pH', 'MEASURE', 'METHOD', 'DOI'.

Store the processed dataframe in the variable `dtm`.

```python
# resulted output
```

Output:
```
ID                               PROTEIN                   SOURCE  \
2          2                             Endolysin  Enterobacteria phage T4   
3          3                             Endolysin  Enterobacteria phage T4   
4          4                             Endolysin  Enterobacteria phage T4   
18        18                             Endolysin  Enterobacteria phage T4   
19        19                             Endolysin  Enterobacteria phage T4   
...      ...                                   ...                      ...   
13332  13331  Cyclin-dependent kinase inhibitor 2A             Homo sapiens   
13333  13332  Cyclin-dependent kinase inhibitor 2A             Homo sapiens   
13334  13333  Cyclin-dependent kinase inhibitor 2A             Homo sapiens   
13335  13334  Cyclin-dependent kinase inhibitor 2A             Homo sapiens   
13336  13335  Cyclin-dependent kinase inhibitor 2A             Homo sapiens   

      UNIPROT           MUTATION MUTATION TYPE  dTm      T   
```

<!-- grade_id: cell-b0925993638de8fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-56ed56a919bd840d, points: 0.5, grade: true, solution: false -->
```python
assert dtm.shape == (3077, 12)
```

<!-- grade_id: cell-81e9048f48dad59e, points: 0.5, grade: true, solution: false -->
```python

```

## Task 2. Formula 1

1 point.

Given the datasets (in the folder `inputs/2/`) with information on the Formula 1 race results for championships from 1950 till the latest 2021 season (9 of 22 races for 2021), compose the ranking of drivers in the latest season and write it to the file `2.out` (rank and driver's surname, tab-separated). 

The dataset `results.csv` contain the results and information about all races. The rest of the datasets contain annotations to identifiers in `results.csv`.

P.S. you may find useful [.join()](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.join.html) function of pandas.

<hr>
2.out:

1	Verstappen\
2	Hamilton\
3	Pérez\
4	Norris\
5	Bottas\
6	Leclerc\
7	Sainz\
8	Ricciardo\
9	Gasly\
10	Vettel\
11	Alonso\
12	Stroll\
13	Ocon\
14	Tsunoda\
15	Räikkönen\
16	Giovinazzi\
17	Russell\
18	Schumacher\
19	Latifi\
20	Mazepin

```python
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-bfddf929eeae93c4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
with open('./inputs/2/circuits.csv') as f:
    Circ = pd.read_csv(f)

with open('./inputs/2/drivers.csv', encoding="utf-8") as f:
    Driv = pd.read_csv(f)

with open('./inputs/2/races.csv') as f:
    Rac = pd.read_csv(f)

with open('./inputs/2/results.csv') as f:
    Res = pd.read_csv(f)
```

```python
ResRac = Res.merge(Rac, on='raceId', how='inner')
ResRac = ResRac[ResRac['year'] >= 2021].sort_values(by='driverId')
ResRac
```

Output:
```
resultId  raceId  driverId  constructorId number  grid position  \
24960     24966    1052         1            131     44     2        1   
25046     25052    1056         1            131     44     7        7   
25081     25087    1059         1            131     44     2        2   
25074     25080    1057         1            131     44     2       15   
25101     25107    1058         1            131     44     2        2   
...         ...     ...       ...            ...    ...   ...      ...   
24995     25001    1053       854            210     47    18       16   
25016     25022    1054       854            210     47    19       17   
25072     25078    1057       854            210     47    17       13   
25037     25043    1055       854            210     47    18       18   
25115     25121    1058       854            210     47    19       16   

      positionText  positionOrder  points  ...  fastestLapTime  \
24960            1              1    25.0  ...      
```

```python
Ranking = pd.DataFrame(columns=['driverId', 'points'])

driverIds = pd.unique(ResRac['driverId'])
points = {}

for id in driverIds:
    points[id] = ResRac[ResRac['driverId'] == id]['points'].sum()

Ranking = pd.DataFrame.from_dict(points, orient='index', columns=['points'])
Ranking['driverId'] = Ranking.index
Ranking = Ranking.merge(Driv, on='driverId', how='inner')
Ranking = Ranking.loc[:, ['points', 'driverId', 'surname']]
Ranking = Ranking.sort_values(by='points', ascending=False)
Ranking
```

Output:
```
points  driverId     surname
7    182.0       830  Verstappen
0    150.0         1    Hamilton
4    104.0       815       Pérez
14   101.0       846      Norris
6     92.0       822      Bottas
13    62.0       844     Leclerc
8     60.0       832       Sainz
5     40.0       817   Ricciardo
12    39.0       842       Gasly
3     30.0        20      Vettel
1     20.0         4      Alonso
10    14.0       840      Stroll
9     12.0       839        Ocon
17     9.0       852     Tsunoda
11     1.0       841  Giovinazzi
2      1.0         8   Räikkönen
15     0.0       847     Russell
16     0.0       849      Latifi
18     0.0       853     Mazepin
19     0.0       854  Schumacher
```

```python
Ranking = Ranking['surname'].reset_index(drop=True)
Ranking.index = Ranking.index + 1
Ranking.to_csv('./2.out', sep='\t', header=False)
```

```python
Ranking
```

Output:
```
7     Verstappen
0       Hamilton
4          Pérez
14        Norris
6         Bottas
13       Leclerc
8          Sainz
5      Ricciardo
12         Gasly
3         Vettel
1         Alonso
10        Stroll
9           Ocon
17       Tsunoda
11    Giovinazzi
2      Räikkönen
15       Russell
16        Latifi
18       Mazepin
19    Schumacher
Name: surname, dtype: object
```

<!-- grade_id: cell-cad795e201495754, points: 1, grade: true, solution: false -->
```python
ans = open('2.out', 'r').read().strip()
example = open('example_outputs/2.out', 'r').read().strip()
assert ans==example
```

## Task 3. Departments

1 point.

Given the table of employees from several departments (`inputs/3.csv`), write the departments from the most populated to the least one with the number of people in it (tab separated) into file `3.out`. In case of equality sort them alphabetically.

<hr>

3.out:

Department of Petroleum Engineering	4\
Analytical Department on Science & Technology Development	3\
Industrial Department	3

<!-- grade_id: cell-fb831a4dc2ec6575, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-7eacf72afa27682a, points: 1, grade: true, solution: false -->
```python
ans = open('3.out', 'r').read().strip()
example = open('example_outputs/3.out', 'r').read().strip()
assert ans==example
```

## Task 4. Shannon entropy

1 point.

Given the table with expression levels of different isoforms for several genes in normal tissues and tumors (`inputs/4.tsv`), calculate the Shannon entropy (H) for each gene for both types of tissues. Calculate the ratio entrory(tumor) / entropy(normal). Write the results into numpy array (6 rows corresponding to each gene, 3 columns corresponding to entropy(normal), entrory(tumor), and their ratio). Round the values in the final array to 2 decimals (use built-in round() function).

Shannon entropy, H:

$$H = - \sum_{i=1}^{n}P(x_{i}) \cdot \log_{2}P(x_{i}) $$

where $x_{i}$ is the expression level of the isoform of the gene, $n$ - number of gene isoforms, $P(x_{i})$ - probability of observing the expression level ($x_{i} / \sum_{i=1}^{n}x_{i}$). 

```python
# output array
```

Output:
```
array([[2.05, 2.32, 1.13],
       [0.7 , 0.98, 1.4 ],
       [1.69, 1.99, 1.17],
       [0.7 , 0.98, 1.4 ],
       [1.69, 1.99, 1.17],
       [1.68, 1.99, 1.18]])
```

<!-- grade_id: cell-8727778d4be05fe9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-10e9e895423456a3, points: 1, grade: true, solution: false -->
```python
assert np.allclose(t4, data)
```

## Task 5. Camera

1 point.

You need to buy a professional camera with an objective. Given a file with the list of all available cameras and objectives that suit your needs, find the cheapest possible combination - assuming that there is only one. Please note that not any pair is compatible.\
The first line in the file contains the space-separated numbers of cameras and objectives. Next goes the information about prices (space-separated as well), and the last set of lines is for the compatibility information. Write a function camera that takes the described file and returns the optimal choice.

Example input file (`input/5.txt`): \
5 3\
C1 22000\
C2 30000\
C3 52000\
C4 25000\
C5 27000\
O1 18000\
O2 20000\
O3 15000\
O1 C2 C4\
O2 C1\
O3 C3 C4 C5

Example output:\
C4 O3 40000


<!-- grade_id: cell-d1a47320f5af85e0, points: 0, grade: false, solution: true -->
```python
def camera(file):
# YOUR CODE HERE

```

<!-- grade_id: cell-a0e60bb7f1caaa19, points: 1, grade: true, solution: false -->
```python
assert camera('inputs/5.txt') == 'C4 O3 40000'
```

## Task 6. Yeast transcriptomics

1 point.

`inputs/6/SC_expression.csv` contains the expression values for each yeast gene over the available conditions. \
`inputs/6/labels_BP.tsv` and `inputs/6/conditions_annotation.csv` contain labels for the individual genes, biological process and conditions of the experiment.

Read the data and select the yeast expression data for those conditions annotated as 'glucose' or 'temperature' and genes with the verified annotation of biological process. Store the resulted dataset with the indexes and columns sorted in alphabetical order in the variable `yeast`.

```python
# resulted output
```

Output:
```
IFFAFA       IFFAFB       IFFAFC       IFFAFF       IFFAFI  \
YAL001C    57.644083    60.021547    66.279060    45.693922    54.536294   
YAL002W    66.519907    59.435399    60.524291    55.726520    68.574339   
YAL003W   798.043808   834.557406   851.110442   861.800145   830.971525   
YAL005C  1101.577458  1227.979888  1227.551689  1149.188468  1198.182568   
YAL007C   109.435976   101.989738   103.188955    96.404144   107.355704   
...              ...          ...          ...          ...          ...   
YPR198W    66.324834    68.930995    70.049426    63.661393    67.564408   
YPR199C    61.740617    64.241812    62.707134    57.824245    57.667081   
YPR200C    10.826554     8.557760     8.532933     8.026078    10.806266   
YPR201W    24.579203    21.101325    23.812836    21.615688    22.824449   
YPR204W     3.998997     3.399658     1.785963     4.377861     2.726815   

              IFFAFQ       IFFAFS      SAABBN      SAABFA      SAABFI  \
YAL001C    47.928108    55.2
```

<!-- grade_id: cell-20fa84329ae9a05a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-48acc1b0019b638b, points: 0.5, grade: true, solution: false -->
```python
assert yeast.shape == (4940, 22)
```

<!-- grade_id: cell-259ce33063010470, points: 0.5, grade: true, solution: false -->
```python
assert yeast.to_dict() == t6.to_dict()
```

## Task 7. Alignment score

2 points

1. Write a function to compute the score of the sequence alignment given the BLOSUM62 matrix (`inputs/7/BLOSUM62.txt`), gap opening penalty (o), and gap extension penalty (e). 

2. Read the files with alignments `7_1.fasta`, `7_2.fasta`, `7_3.fasta`, compute the scores for each alignment (gap opening penalty = -10, and gap extension penalty = -1) and write the results into the dataframe `data` with the columns: 

`id_1`, `id_2` - ids of the aligned sequences <br>
`protein_1`, `protein_2` - names of the aligned proteins<br>
`organism_1`, `organism_2` - organisms <br>
`score` - computed score for the alignment

P.S. the sequence header has the following format `>1WLA_1|Chain A|MYOGLOBIN|Equus caballus (9796)`, where `|` separates the id, chain, protein name, and organism.

```python
# output dataframe
```

Output:
```
id_1   protein_1      organism_1    id_2    protein_2        organism_2  \
0  1WLA_1   MYOGLOBIN  Equus caballus  1OJ6_1  Neuroglobin      Homo sapiens   
1  5DUV_1  Galectin-4    Homo sapiens  4O4Z_1  Neuroglobin      Mus musculus   
2  1MNK_1   MYOGLOBIN      Sus scrofa  2SPM_1    MYOGLOBIN  Physeter catodon   

   score  
0     19  
1    -63  
2    663
```

<!-- grade_id: cell-3336821f5dc7da2c, points: 0, grade: false, solution: true -->
```python
def score(seq1,seq2, o, e):
# YOUR CODE HERE

```

<!-- grade_id: cell-7d97bbd91c74ae9e, points: 1, grade: true, solution: false -->
```python
# CHECK FUNCTION
seq1 = 'MVLSEGEWQLVLHVWAKVEADVAGHGQDIFIRLFKSHPETLEKFD-RFKHLKTEAEMKASEDLKKHGVTVLTALGAILKKKGHHEAELKPLAQSHATKHK-IPIKYLEF--ISEAIIHVLHSRHPGNFGADAQGAMNKALELFRKDIAAKYKELGYQG-'
seq2 = 'MERPEPE--LIRQSWRAVSRSPLEHGTVLFARLFALEPDLLPLFQYNGRQFSSPEDSLSSPEFLDHIRKVMLVIDAAVTNV-EDLSSLEEYLASLGRKHRAVGVKLSSFSTVGESLLYMLEKSLGPAFTPATRAAWS---QLYGAVVQAMSR--GWDGE'
assert score(seq1,seq2, -10, -1) == 47
```

<!-- grade_id: cell-2a2ee445ecd97abb, points: 1, grade: true, solution: false -->
```python
# CHECK DATASET
assert data.to_dict() == t7.to_dict()
```

## Task 8. Dihedral angles

1 point.

Given the table with backbone dihedral angles $\phi$ and $\psi$ in radians, convert them into degrees and round to 2 decimals (use built-in round function). Calculate the fraction of residues involved in alpha-helix formation given the average dihedral angles $\phi$ and $\psi$ for alpha-helix (-64 +/- 7, -41 +/- 7). Store the array and the fraction in the variables `phipsi` and `alpha`, respectively.

<!-- grade_id: cell-634b9530024a7018, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-23b67b4fc926a691, points: 1, grade: true, solution: false -->
```python
assert alpha == 0.17
assert np.allclose(phipsi, t8, equal_nan=True)
```